Osnabrück University - Machine Learning (Summer Term 2017) - Prof. Dr.-Ing. G. Heidemann, Ulf Krumnack

# Exercise Sheet 02: Decision Trees

## Introduction
By now everyone should have found a group. If someone still has none but wants to participate in the course please contact one of the tutors.

This week's sheet should be solved and handed in before the end of **Sunday, April 22, 2018**. If you need help (and Google and other resources were not enough), feel free to contact your groups designated tutor or whom ever of us you run into first. Please upload your results to your group's studip folder.

## Assignment 0: Math recap (Euclidean Space) [2 Bonus Points]

This exercise is supposed to be very easy and is voluntary. There will be a similar exercise on every sheet.
It is intended to revise some basic mathematical notions that are assumed throughout this class and to allow you to check if you are comfortable with them.
Usually you should have no problem to answer these questions offhand, but if you feel unsure, this is a good time to look them up again. You are always welcome to discuss questions with the tutors or in the practice session.
Also, if you have a (math) topic you would like to recap, please let us know.

**a)** What is a *Euclidean space*? What is the *Cartesian plane*? How are they usually denoted? How to write points in these spaces?

<!-- grade_id: math-euclid-a1, points: 0, grade: true, solution: true -->
A *Euclidean space* is a space in which the axioms of Euclidean geometry hold. There is one such space for each dimension $n$. One usually uses Cartesian coordinates, i.e. tuples of real numbers, to refer to points in these spaces, so the *Cartesian plane* is the two-dimensional Eucliden space modeled by pairs of real coordinates. The $n$-dimensional Euclidean space is usually denoted by
$$\mathbb{R}^n = \underbrace{\mathbb{R}\times\ldots\times\mathbb{R}}_{n\text{-times}}$$
(sometimes also $\mathbf{R}^n$). Points in that space are hence denoted by $n$-tuples of real numbers,
$$\vec{x}=\mathbf{x} = (x_1,\ldots,x_n)\in\mathbb{R}^n$$

**b)** What is the *norm* of a vector in a Euclidean space? How to *add* and *substract* two vectors? How is the *Euclidean distance* defined? Are there other ways to measure distances?

<!-- grade_id: math-euclid-a2, points: 0, grade: true, solution: true -->
The *(Euclidean) norm* of a vector

$$\|\vec{x}\| = \sqrt{x_1^2+\ldots+x_n^2}$$

it can be interpreted as the *length* of the vector. The sum of two vectors is defined as
$$\vec{x}+\vec{y} = \begin{pmatrix} x_1 \\ \vdots \\ x_n\end{pmatrix}+
\begin{pmatrix} y_1 \\ \vdots \\ y_n\end{pmatrix} =
\begin{pmatrix} x_1+y_1 \\ \vdots \\ x_n+y_n\end{pmatrix}$$
and analogously the difference is
$$\vec{x}-\vec{y} = \begin{pmatrix} x_1 \\ \vdots \\ x_n\end{pmatrix}-
\begin{pmatrix} y_1 \\ \vdots \\ y_n\end{pmatrix} =
\begin{pmatrix} x_1-y_1 \\ \vdots \\ x_n-y_n\end{pmatrix}$$
The euclidean distance is defined as the norm of the difference of two vectors

$$d(\vec{x},\vec{y}) = \|\vec{x}-\vec{y}\| = \sqrt{(x_1-y_1)^2+\ldots+(x_n-y_n)^2} = \left(\sum_{i=1}^n (x_i-y_i)^{2}\right)^{\frac{1}{2}}$$

(remember the Pythagorean theorem). Be replacing 2 by $p$ one gets the $p$-norm instead of the Euclidean norm, from which one then gets the $p$-metric to measure distances. There are even more distance measures.

**c)** What is the (standard) *scalar product* of two vectors? How is it related to the length and angle between these vectors? Name some use cases.

<!-- grade_id: math-euclid-a3, points: 0, grade: true, solution: true -->
The scalar product of two vectors $\vec{x}$ and $\vec{y}$ is defined as

$$ \vec{x}\cdot\vec{y} = \langle\vec{x},\vec{y}\rangle = x_1y_1+\ldots+x_ny_n = \sum_{i=1}^{n}x_iy_i$$

On a geometric level the scalar product can be computed as

$$\vec{x}\cdot\vec{y} = \|\vec{x}\| \|\vec{y}\| \cos(\sphericalangle(\vec{x},\vec{y}))$$

where $\sphericalangle(\vec{x},\vec{y})$ denotes the angle between vectors $\vec{x}$ and $\vec{y}$. The scalar product is sometimes used as a similarity measure, which will be $0$ for orthogonal (i.e. unrelated) vectors.

## Assignment 1: Decision Trees [2 Points]
Draw the decision trees for the following boolean functions. Either use pen and paper and bring the results to the feedback session or employ your ASCII artist within below. 

Note: $\oplus := xor$, that means one of the operands has to be true, while the other one has to be false:

$\oplus$ | $B$ | $\neg B$
---------|-----|---------
$A$      |  f  |    t
$\neg A$ |  t  |    f

### a) $\neg A \wedge B$

<!-- grade_id: 1a_answer, points: 0.5, grade: true, solution: true -->

         A
        / \
     t /   \ f
      /     \
     No      B
            / \
         t /   \ f
          /     \
        Yes     No


### b) $A \oplus B$ 

<!-- grade_id: 1b_answer, points: 0.5, grade: true, solution: true -->
               A
              / \
         t  /     \  f
          /         \
         B           B
        / \         / \
     t /   \ f   t /   \ f
      /     \     /     \
     No     Yes Yes     No

### c) $A \vee (B \wedge C) \vee (\neg C \wedge D)$

<!-- grade_id: 1c_answer, points: 0.5, grade: true, solution: true -->
            A
           / \
      t  /     \  f
       /         \
     Yes          C
                 / \
            t  /     \  f
             /         \
            B           D
           / \         / \
        t /   \ f   t /   \ f
         /     \     /     \
       Yes     No  Yes     No

### d) $(A \rightarrow (B \wedge \neg C)) \vee (A \wedge B)$

<!-- grade_id: 1d_answer, points: 0.5, grade: true, solution: true -->
$= (\neg A \vee (B \wedge \neg C)) \vee (A \wedge B)$

$= \neg A \vee (A \wedge B) \vee (B \wedge \neg C)$

$= \neg A \vee (A \wedge B)$

$= \neg A \vee  B$

              A
             / \
          t /   \ f
           /     \
          B      Yes
         / \
      t /   \ f
       /     \
     Yes     No

## Assignment 2: Entropy and Information Gain [8 Points]

In many machine learning applications it is crucial to determine which criterions are necessary for a good classification. Decision trees have those criterions close to the root, imposing an order from significant to less significant criterions. One way to select the most important criterion is to compare its information gain or its entropy to others. The following dataset is a hands-on example for this method. 

Consider the following attributes with their possible values:

  * $raining = \{yes, no\}$
  * $tired = \{yes, no\}$
  * $late = \{yes, no\}$
  * $distance = \{short, medium, long\}$

And a training data set consisting of those attributes:

| #  | raining | tired | late | distance | attend_party |
|----|---------|-------|------|----------|--------------|
| 1  | yes     | no    | no   | short    | **yes**      |
| 2  | yes     | no    | yes  | medium   | **no**       |
| 3  | no      | yes   | no   | long     | **no**       |
| 4  | yes     | yes   | yes  | short    | **no**       |
| 5  | yes     | no    | no   | short    | **yes**      |
| 6  | no      | no    | no   | medium   | **yes**      |
| 7  | no      | yes   | no   | long     | **no**       |
| 8  | yes     | no    | yes  | short    | **no**       |
| 9  | yes     | yes   | no   | short    | **yes**      |
| 10 | no      | yes   | no   | medium   | **no**       |
| 11 | no      | yes   | no   | long     | **no**       |
| 12 | no      | yes   | yes  | short    | **no**       |

### a)

Build the root node of a decision tree from the training samples given in the table above by calculating the information gain for all four attributes (raining, tired, late, distance).

$$\operatorname{Gain}(S,A) = \operatorname{Entropy}(S) - \sum_{v \in \operatorname{Values}(A)} \frac{|S_v|}{|S|}\operatorname{Entropy}(S_v)$$

$$\operatorname{Entropy}(S) = -p_{\oplus} log_{2} p_{\oplus} - p_{\ominus} log_{2} p_{\ominus}$$

$S$ is the set of all data samples. $S_v$ is the subset for which attribute $A$ has value $v$. An example for attribute **tired** with value $yes$ would be:
$$|S_{yes}| = 7, S_{yes}:[1+, 6−]$$

<!-- grade_id: 2a_answer, points: 4, grade: true, solution: true -->
*Caution: some pretty intense rounding was applied to the following numbers. So your results might differ, but they should be in the same ballpark!*

Entropy of the whole dataset:

$$Entropy\left(S\right) = -\frac{4}{12} \log_{2} \frac{4}{12} - \frac{8}{12} \log_{2} \frac{8}{12} \approx 0.92$$

Attribute **raining**: 
$$\left|S_{yes}\right|:6 , S_{yes}:[3+,3-]$$

$$Entropy\left(S_{yes}\right) = -\frac{3}{6} \log_{2} \frac{3}{6}-\frac{3}{6} \log_{2} \frac{3}{6} = 1$$

$$\left|S_{no}\right|:6 , S_{no}:[1+,5-]$$

$$Entropy\left(S_{no}\right) = -\frac{1}{6} \log_{2} \frac{1}{6}-\frac{5}{6} \log_{2} \frac{5}{6} \approx 0.65$$

$$Gain\left(S,raining\right) \approx 0.92 - \left(\frac{6}{12}\cdot 1 + \frac{6}{12}\cdot 0.65\right) = 0.095$$

Attribute **tired**:
$$\left|S_{yes}\right|:7 , S_{yes}:[1+,6-]$$

$$Entropy\left(S_{yes}\right) = -\frac{1}{7} \log_{2} \frac{1}{7}-\frac{6}{7} \log_{2} \frac{6}{7} \approx 0.59$$

$$\left|S_{no}\right|:5 , S_{no}:[3+,2-]$$

$$Entropy\left(S_{no}\right) = -\frac{3}{5} \log_{2} \frac{3}{5}-\frac{2}{5} \log_{2} \frac{2}{5} \approx 0.97$$

$$Gain\left(S,tired\right) \approx 0.92 - \left(\frac{7}{12}\cdot 0.59 + \frac{5}{12}\cdot 0.97\right) \approx 0.171$$

Attribute **late**:
$$\left|S_{yes}\right|:4 , S_{yes}:[0+,4-]$$

$$Entropy\left(S_{yes}\right) = -\frac{0}{4} \log_{2} \frac{0}{4}-\frac{4}{4} \log_{2} \frac{4}{4} = 0$$

$$\left|S_{no}\right|:8 , S_{no}:[4+,4-]$$

$$Entropy\left(S_{no}\right) = -\frac{4}{8} \log_{2} \frac{4}{8}-\frac{4}{8} \log_{2} \frac{4}{8} = 1$$

$$Gain\left(S,late\right) \approx 0.92 - \left(\frac{4}{12}\cdot 0 + \frac{8}{12}\cdot 1\right) \approx 0.253$$

Attribute **distance**:
$$\left|S_{short}\right|:6 , S_{short}:[3+,3-]$$

$$Entropy\left(S_{short}\right) = -\frac{3}{6} \log_{2} \frac{3}{6}-\frac{3}{6} \log_{2} \frac{3}{6} = 1$$

$$\left|S_{medium}\right|:3 , S_{medium}:[1+,2-]$$

$$Entropy\left(S_{medium}\right) = -\frac{1}{3} \log_{2} \frac{1}{3}-\frac{2}{3} \log_{2} \frac{2}{3} \approx 0.918$$

$$\left|S_{long}\right|:3 , S_{long}:[0+,3-]$$

$$Entropy\left(S_{long}\right) = -\frac{0}{3} \log_{2} \frac{0}{3}-\frac{3}{3} \log_{2} \frac{3}{3} = 0$$

$$Gain\left(S,distance\right) \approx 0.92 - \left(\frac{6}{12}\cdot 1 + \frac{3}{12}\cdot 0.918 + \frac{3}{12}\cdot 0\right) = 0.191$$

The information gain is greatest for the **late** attribute. 

### b)

Perform the same calculation as in **a)** but use the gain ratio instead of the information gain. Does the result for the root node change?

$$\operatorname{GainRatio}(S,A) = \frac{\operatorname{Gain}(S,A)}{\operatorname{SplitInformation}(S,A)}$$

$$\operatorname{SplitInformation}(S,A) = - \sum_{v \in \operatorname{Values}(A)} \frac{|S_v|}{|S|} \log_{2} \frac{|S_{v}|}{|S|}$$

<!-- grade_id: 2b_answer, points: 4, grade: true, solution: true -->
Attribute **raining**:

$$SplitInformation(S,raining) = -\left(\frac{6}{12} \log_{2} \frac{6}{12} + \frac{6}{12} \log_{2} \frac{6}{12}\right) = 1$$
$$GainRatio(S,raining) = \frac{0.095}{1} = 0.095$$

Attribute **tired**:

$$SplitInformation(S,tired) = -\left(\frac{7}{12} \log_{2} \frac{7}{12} + \frac{5}{12} \log_{2} \frac{5}{12}\right) = 0.98$$
$$GainRatio(S, tired) = \frac{0.171}{0.98} \approx 0.174$$

Attribute **late**:

$$SplitInformation(S,late) = -\left(\frac{4}{12} \log_{2} \frac{4}{12} + \frac{8}{12} \log_{2} \frac{8}{12}\right) = 0.918$$
$$GainRatio(S,late) = \frac{0.253}{0.918} \approx 0.276$$

Attribute **distance**:

$$SplitInformation(S,distance) = -\left(\frac{6}{12} \log_{2} \frac{6}{12} + \frac{3}{12} \log_{2} \frac{3}{12} + \frac{3}{12} \log_{2} \frac{3}{12}\right) = 1.5$$
$$GainRatio(S,distance) = \frac{0.191}{1.5} \approx 0.127$$

We should still use the **late** attribute.

## Assignment 3: ID3 algorithm [5 Points]

Implement the following two functions in Python. Take a look at the `assert`s to see how the function should behave. An assert is a condition that your function is required to pass. Most of the conditions here are taken from the lecture slides (ML-03, Slide 12 & 13). Don't worry if you do not get all asserts to pass, just comment the failing ones out.

### a) Entropy

$$\operatorname{Entropy}(S) = - \sum_{i=1...c} p_i \log_2 p_i$$

<!-- grade_id: 3a_code, points: 0, grade: false, solution: true -->
```python
from math import log2
def entropy(s):
    """
    Calculate the entropy for a given target value set. 
    
    Args:
        s (list): Target classes for specific observations.
        
    Returns:
        The entropy of s.
    """
    ### BEGIN SOLUTION
    return -sum((s.count(target) / len(s)) * log2(s.count(target) / len(s)) 
                for target in set(s))
    ### END SOLUTION

# See ML-03, Slide 12 & 13
```

<!-- grade_id: 3a_test, points: 1, grade: true, solution: false -->
```python
assert entropy([1,1,1,0,0,0]) == 1.0
assert round(entropy([1,1,1,1,0,0,0]), 3) == 0.985
assert round(entropy([1,1,1,1,1,1,0]), 3) == 0.592
assert round(entropy([1,1,1,1,1,1,0,0]), 3) == 0.811
assert round(entropy([2,2,1,1,0,0]), 3) == 1.585
assert round(entropy([2,2,2,1,0]), 3) == 1.371
assert round(entropy([2,2,2,0,0]), 3) == 0.971
assert entropy(['yes','yes','yes','no','no','no']) == 1.0
```

### b)  Information Gain

$$\operatorname{Gain}(S,A) = \operatorname{Entropy}(S) - \sum_{v \in \operatorname{Values}(A)} \frac{|S_v|}{|S|} \operatorname{Entropy}(S_v)$$

<!-- grade_id: 3b_code, points: 0, grade: false, solution: true -->
```python
def gain(targets, attr_values):
    """
    Calculates the expected reduction in entropy due to sorting on A.
    
    Args:
        targets (list): Target classes for observations in attr_values.
        attr_values (list): Values of each instance for the respective attribute.
        
    Returns:
        The information gain of 
    """
    ### BEGIN SOLUTION
    sigma = 0
    for v in set(attr_values): # Sets only contain unique values.
        S_v = [targets[key] for (key, v_) in enumerate(attr_values) if v_ == v]
        sigma += ((len(S_v) / len(targets)) * entropy(S_v))
    return entropy(targets) - sigma
    ### END SOLUTION

# See ML-03, Slide 12 & 13

```

<!-- grade_id: 3b_test, points: 1, grade: true, solution: false -->
```python
# The lists here can each be seen as one column of a table such as the one in assignment 2. 
# Assert targets would be the last column, while the attribute values are the values of one attribute, here the 
# example rain and distance
assert_targets = ["no","no","yes","yes","yes","no","yes","no","yes","yes","yes","yes","yes","no"]
assert_attribute_values_1 = ["yes", "yes","yes","yes","no", "no", "no", "yes", "no", "no", "no","yes", "no", "yes"]
assert_attribute_values_2 = ["high","low","medium","high","high","medium","low","medium","low","high","high","medium","low","low"]
assert_attribute_values_3 = [0,1,0,0,0,1,1,0,0,0,1,1,0,1]
assert round(gain(assert_targets, assert_attribute_values_1), 3) == 0.152
assert round(gain(assert_targets, assert_attribute_values_2), 3) == 0.05
assert round(gain(assert_targets, assert_attribute_values_3), 3) == 0.048
```

### c) ID3

In the next two cells we have implemented the ID3 algorithm following the pseudocode from [Wikipedia](https://en.wikipedia.org/wiki/ID3_algorithm#Pseudocode) - it relies on your two functions from above, `entropy` and `gain`. Try to understand what the code does and replace `YOUR CODE HERE` with meaningful comments describing the respective parts of the code. Though its often annoying, being able to read other peoples code is one of the key skills (and obstacle) in software engineering. So give it a try! Otherwise you are of course welcome to write your own implementation.

<!-- grade_id: 3c_answer, points: 3, grade: true, solution: true -->
```python
from collections import Counter, namedtuple


class Node(namedtuple('Node', 'label children')):
    """
    A small node representation with a pretty string representation.
    """
    def __str__(self, level=0):
        return_str ='{}{!s}\n'.format(' ' * level * 4, self.label)
        for child in self.children:
            return_str += child.__str__(level + 1)
        return return_str

def id3(examples, attributes, target_attribute=None): 
    """
    Calculate a tree of Nodes (fields: label [string], children [list]) 
    using the ID3 algorithm found as pseudocode on Wikipedia.
    """
    ### BEGIN SOLUTION
    # Create a root node for the tree
    # If all examples are positive, return the single-node tree Root, with label = 
    # If all examples are negative, return the single-node tree Root, with label = -.
    ### END SOLUTION
    if all(target == examples['targets'][0] for target in examples['targets']):
        return Node('Result: {!s}'.format(examples['target_names'][examples['targets'][0]]), [])
    
    ### BEGIN SOLUTION
    # If number of predicting attributes is empty, then return the single node tree Root,
    # with label = most common value of the target attribute in the examples.
    ### END SOLUTION
    if len(attributes) == 0:
        attr = Counter(data_sample[target_attribute] for data_sample in examples['data']).most_common(1)
        return Node('Attribute: {!s}, {!s} occurences'.format(examples['attributes'][target_attribute], attr), [])
    
    ### BEGIN SOLUTION
    # A <-- The Attribute that best classifies examples.
    ### END SOLUTION
    gains = [gain(examples['targets'], [r[attribute] for r in examples['data']]) 
             for attribute in attributes]
    max_gain_attribute = attributes[gains.index(max(gains))]
    
    ### BEGIN SOLUTION
    # Create a root node for the tree (2)
    # Decision Tree attribute for Root = A.
    ### END SOLUTION
    root = Node('Attribute: {!s} (gain {!s})'.format(examples['attributes'][max_gain_attribute], 
                                                     round(max(gains), 4)), [])
    
    ### BEGIN SOLUTION
    # Otherwise Begin
    # For each possible value, vi, of A,
    ### END SOLUTION
    for vi in set(data_sample[max_gain_attribute] for data_sample in examples['data']):
        ### BEGIN SOLUTION
        # Add a new tree branch below Root, corresponding to the test A = vi.
        ### END SOLUTION
        child = Node('Value: {!s}'.format(vi), [])
        root.children.append(child)
        
        ### BEGIN SOLUTION
        # Let Examples(vi) be the subset of examples that have the value vi for A.
        ### END SOLUTION
        vi_indices = [idx for idx, data_sample in enumerate(examples['data']) 
                          if data_sample[max_gain_attribute] == vi]
        examples_vi = dict(examples)
        examples_vi['data'] = [examples['data'][i] for i in vi_indices]
        examples_vi['targets'] = [examples['targets'][i] for i in vi_indices]
        
        if examples_vi['data']:
            ### BEGIN SOLUTION
            # If Examples(vi) is not empty below this new branch 
            # add the subtree ID3 (Examples(vi), Attributes – {A}).
            ### END SOLUTION
            child.children.append(
                id3(examples_vi,
                    [attribute_ for attribute_ in attributes if not attribute_ == max_gain_attribute],
                    max_gain_attribute)
            )
            
        else:
            ### BEGIN SOLUTION
            # If Examples(vi) is empty
            # Then below this new branch add a leaf node with label = most common target value in the examples.
            ### END SOLUTION
            attr = Counter(examples_vi['targets']).most_common(1)
            label = 'Attribute: {!s}, {!s} occurences'.format(examples['attributes'][target_attribute], attr)
            child.children.append(Node(label, []))

    return root
```

### d) The algorithm is applied to two data sets. Run those and discuss the differences. For which data set is the ID3 algorithm better suited and why?

YOUR ANSWER HERE

The problem with the iris data set is that since the ID3 algorithm works with the assumption of nominal variables it splits on every unique value for continuous data. Thus the tree grows very wide.

This code runs the ID3 algorithm on the party data set which you already know from assignment 2.

```python
import json

with open('party.json', 'r') as party_file:
    party = json.load(party_file)

# Make sure our gain function handles the data set as expected.
assert round(gain(party['targets'], [r[2] for r in party['data']]), 3) == 0.252

# Apply ID3 algorithm
tree_party = id3(party, list(range(len(party['attributes']))))

print(tree_party)
```

Output:
```
Attribute: late (gain 0.2516)
    Value: yes
        Result: no
    Value: no
        Attribute: distance (gain 0.75)
            Value: long
                Result: no
            Value: short
                Result: yes
            Value: medium
                Attribute: tired (gain 1.0)
                    Value: yes
                        Result: no
                    Value: no
                        Result: yes
```

This code runs the ID3 algorithm on the famous iris flowser data set, which you will hear more about in assignment 4 below.

```python
import json

with open('iris.json', 'r') as iris_file:
    iris = json.load(iris_file)

# Make sure our gain function handles the data set as expected.
assert round(gain(iris['targets'], [r[2] for r in iris['data']]), 3) == 1.446

# Apply ID3 algorithm
tree_iris = id3(iris, list(range(len(iris['attributes']))))

print(tree_iris)
```

Output:
```
Attribute: petal length (gain 1.4463)
    Value: 3.3
        Result: Iris-versicolor
    Value: 4.2
        Result: Iris-versicolor
    Value: 4.7
        Result: Iris-versicolor
    Value: 6.1
        Result: Iris-virginica
    Value: 1.6
        Result: Iris-setosa
    Value: 5.7
        Result: Iris-virginica
    Value: 5.2
        Result: Iris-virginica
    Value: 1.0
        Result: Iris-setosa
    Value: 3.0
        Result: Iris-versicolor
    Value: 5.6
        Result: Iris-virginica
    Value: 5.9
        Result: Iris-virginica
    Value: 5.3
        Result: Iris-virginica
    Value: 4.1
        Result: Iris-versicolor
    Value: 3.6
        Result: Iris-versicolor
    Value: 5.1
        Attribute: sepal length (gain 0.5436)
            Value: 5.8
                Result: Iris-virginica
            Value: 6.5
                Result: Iris-virginica
            Value: 6.0
                Result: Iris-versicolor
            Value: 6.3
                Result: Iris-virginica
        
```

## Assignment 4: Decision Trees on Iris Flowers [5 Points]

In this exercise we are going to examine and compare two decision trees that were generated from the iris flower data set ([Wikipedia](https://en.wikipedia.org/wiki/Iris_flower_data_set)) where three variations of the iris flower are quantified. The Iris data set is a classical example of a labeled dataset, i.e. every sample consists of two parts: features and labels. There are four features per sample in this data set (sepal length ($x_1$), sepal width ($x_2$), petal length ($x_2$) and petal width ($x_4$) in cm) and a corresponding label (Iris Setosa, Iris Versicolour, Iris Virginica). These samples are by nature **noisy**, no matter how carefully the measurement was taken - slight deviation from the actual length **cannot be avoided**. We want to learn how the features are related to the label so that we could (in the future) predict the label of a new sample automatically. One way to obtain such a `classifier` is to train a decision tree on the data.

Here are two decisions tree generated by the data set. We will now take a closer look.

### Tree 1:

                      +
                      |
                      |
                      |
       x3 < 2.45      |     x3 >= 2.45
   +------------------+------------------+
   |                                     |
   |                        x4 < 1.75    |     x4 >= 1.75
   +                           +---------+---------+
setosa                         |                   |
                               |                   |
                     x3 < 4.95 |   x3 >= 4.95      +
                        +--------------+       virginica
                        |              |
                        |              |
              x4 < 1.65 | x4 >= 1.65   +
                 +------------+    virginica
                 |            |
                 |            |
                 +            +
            versicolor    virginica


### Tree 2:

                      +
                      |
                      |
                      |
       x3 < 2.45      |     x3 >= 2.45
   +------------------+------------------+
   |                                     |
   |                        x4 < 1.75    |     x4 >= 1.75
   +                           +---------+---------+
setosa                         |                   |
                               |                   |
                               +                   +
                          versicolor           virginica


### a)

What does it mean that the features $x1$ and $x2$ do not appear in the decision trees?

<!-- grade_id: 4a, points: 1, grade: true, solution: true -->
Sepal length and sepal width are not relevant for the classification. This might be either because they are redundant or because they are independent of the class.

### b)
With which method from the lecture might the second tree have been generated from the first one? Explain the procedure.

<!-- grade_id: 4b, points: 1, grade: true, solution: true -->
Reduced error pruning. Greedily remove the node that reduces error on validation set the most.

### c)
After training the tree we can calculate the accuracy, i.e. the percentage of the training set that is classified correctly. Although the first tree was trained on the data set until no improvement of the accuracy was possible, its accuracy is *only* 98%. Explain why it is not 100 %

<!-- grade_id: 4c, points: 1, grade: true, solution: true -->
The dataset is probably inconsistent, i.e. there are samples with the same features but different classes. Alternative: the thresholding has not produced the optimal partitioning.

### d)

Tree 2 only has a 96% accuracy on the training set. Why might this tree still be preferable over tree 1?

<!-- grade_id: 4d, points: 1, grade: true, solution: true -->
Tree 1 is probably overfitted to this specific dataset, i.e. it has not only captured the structure but also the noise in the data. It probably won't generalize as well as the second tree.
Another advantage of tree 2 is that it is faster at classifiying new data since less computations have to be made. This difference is hardly noticeable however.